# Module V – Auto-Graded Lab: Data Mining, Visualization & Modeling

**Instructions:**
- Do not modify locked cells
- Write code only where indicated
- Run all cells before submission
- Dataset: Modeling_Customer_Churn_Module_V.csv


## Question 1
**Load the dataset and perform a train-test split**

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import pandas as pd
from sklearn.model_selection import train_test_split

df = pd.read_csv('Modeling_Customer_Churn_Module_V.csv')
X = df.drop(columns=['Churn'])
y = df['Churn']
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.3, random_state=42)
X_train.shape, X_test.shape
```

Output:
```
((280, 5), (120, 5))
```

<!-- grade_id: None, points: 5, grade: true, solution: false -->
```python
assert X_train.shape[0] > X_test.shape[0]
```

## Question 2
**Train a Logistic Regression model and report accuracy**

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
from sklearn.linear_model import LogisticRegression
from sklearn.metrics import accuracy_score

model = LogisticRegression(max_iter=1000)
model.fit(X_train, y_train)
y_pred = model.predict(X_test)
acc = accuracy_score(y_test, y_pred)
acc
```

Output:
```
0.675
```

<!-- grade_id: None, points: 10, grade: true, solution: false -->
```python
assert acc >= 0.5
```

## Question 3
**Evaluate the model using precision, recall, and F1-score**

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
from sklearn.metrics import classification_report

report = classification_report(y_test, y_pred, output_dict=True)
report
```

<!-- grade_id: None, points: 10, grade: true, solution: false -->
```python
assert '1' in report
```

## Question 4
**Train a Decision Tree and compare accuracy**

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
from sklearn.tree import DecisionTreeClassifier

dt = DecisionTreeClassifier(random_state=42)
dt.fit(X_train, y_train)
dt_pred = dt.predict(X_test)
dt_acc = accuracy_score(y_test, dt_pred)
dt_acc
```

<!-- grade_id: None, points: 10, grade: true, solution: false -->
```python
assert dt_acc >= 0.4
```