# Exploratory Data Analysis (EDA)

**Instructions:**
- Do not modify locked cells
- Write code only where indicated
- Run all cells before submission
- Dataset: EDA_Customer_Behavior_Module_IV.csv


## Question 1
**Load the dataset and compute basic descriptive statistics**

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import pandas as pd
df = pd.read_csv('EDA_Customer_Behavior_Module_IV.csv')
df.describe()
```

<!-- grade_id: None, points: 5, grade: true, solution: false -->
```python
assert df.shape[0] >= 200
```

## Question 2
**Create a histogram for Purchase_Amount**

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import matplotlib.pyplot as plt
df['Purchase_Amount'].hist()
plt.xlabel('Purchase Amount')
plt.ylabel('Frequency')
plt.show()
```

<!-- grade_id: None, points: 5, grade: true, solution: false -->
```python
assert 'Purchase_Amount' in df.columns
```

## Question 3
**Compute the correlation matrix and identify the strongest correlation**

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
corr = df.corr(numeric_only=True)
corr
```

<!-- grade_id: None, points: 10, grade: true, solution: false -->
```python
assert corr.shape[0] >= 3
```

## Question 4
**Scatterplot between Annual_Income and Spending_Score**

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
plt.scatter(df['Annual_Income'], df['Spending_Score'])
plt.xlabel('Annual Income')
plt.ylabel('Spending Score')
plt.show()
```

<!-- grade_id: None, points: 10, grade: true, solution: false -->
```python
assert 'Annual_Income' in df.columns and 'Spending_Score' in df.columns
```