# Module II – Auto-Graded Lab: Python & R for Data Analytics

**Instructions:**
- Do not modify locked cells
- Write code only where indicated
- Run all cells before submission
- Dataset: Customer_Analytics_Module_II.csv


## Question 1
**Load the dataset and display basic information**

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import pandas as pd
df = pd.read_csv('Customer_Analytics_Module_II.csv')
df.info()
```

<!-- grade_id: None, points: 5, grade: true, solution: false -->
```python
assert df.shape[0] >= 100
```

## Question 2
**Compute average age and annual income**

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
avg_age = df['Age'].mean()
avg_income = df['Annual_Income'].mean()
avg_age, avg_income
```

<!-- grade_id: None, points: 5, grade: true, solution: false -->
```python
assert avg_age > 18 and avg_income > 20000
```

## Question 3
**Group by Gender and compute churn rate**

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
churn_rate = df.groupby('Gender')['Churn'].mean()
churn_rate
```

<!-- grade_id: None, points: 10, grade: true, solution: false -->
```python
assert churn_rate.shape[0] >= 2
```

## Question 4
**Create a bar plot of churn rate by gender**

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import matplotlib.pyplot as plt
churn_rate.plot(kind='bar')
plt.title('Churn Rate by Gender')
plt.show()
```

<!-- grade_id: None, points: 10, grade: true, solution: false -->
```python
assert 'churn_rate' in globals()
```

## Question 5 
**Discuss reproducibility challenges when performing analytics in Python vs R.**

Write a concise, research-oriented response.

<!-- grade_id: None, points: 10, grade: true, solution: false -->
YOUR ANSWER HERE